```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

# Lab 7: A/B Testing Application

Welcome to Data 111 Lab 7! This week's lab will focus on an interesting application of A/B Testing. We will begin with a discussion and then do an exercise using data from Karuk and Yurok tribes. Since we are already familiar with A/B testing, let's focus our attention on the discussion today. This discussion may include some heavy topics, so take your time to ponder and share throughout this process. 

**Submission**: Once you’re finished, run all cells besides the last one, select File > Save Notebook, and then execute the final cell. Then submit the downloaded zip file, that includes your notebook..

First, set up the notebook by running the cell below.

```python
# Run this cell to set up the notebook, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', (FutureWarning, np.VisibleDeprecationWarning))
```

# Discussion: Traditional Ecological Knowledge

Indigenous people are the people who have inhabited a land area from the earliest times, before colonization, and have a way of life that is *inextricably linked* to that land. Over many thousands of years and more, indegenous folks learn and pass down wisdom about how humans can best support the local ecosystem they are a part of and depend on. This wisdom is sometimes called "Traditional Ecological Knowledge" (TEK). Invasion by colonization across the world has caused many genocides of indigenous people and their lands, and has attempted to erase indigenous culture and TEK. Watch [this video](https://www.youtube.com/watch?v=T-azcPugmKQ) to see the impacts of colonialism on indigenous peoples here and discuss the following question:

<!-- BEGIN QUESTION -->

**Discussion Question 1:** What are the ways that colonialism harmed/harms the indigenous people of our region?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Discussion Question 2:** "Western" science (like you are learning in this class) has been dominated by folks from colonialist cultures. It is a very young culture compared to indigenous culture, and it is perhaps not surprising that recent scientific studies are relearning the lessons of the much older indigenous wisdom. Watch [this video](https://www.youtube.com/watch?v=wlzDJQjjQDg) about indigenous cultural burning and wildfire management to answer the following question. 

How do prescribed burns help with wildfire management?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Discussion Question 3:** In addition to protecting people and the ecosystems we are part of from dangerous wildfires, perscribed burns are needed for cultural resources that are essential to the Karuk and Yurok way of life. For this reason, Karuk and Yurok people (and many others indigenous peoples) refer to their prescribed fire as *cultural burns* that are central to subsistence and ceremonial practices. The term *cultural burning* distinguishes these fires from the fuel reduction-focused *prescribed burns* of public land agencies whose primary objective is to reduce fuel loads and moderate wildfire intensity. One of the ways cultural burns are central to Karuk and Yurok way of life, is that they provide more hazelnut stems that are needed for *basket weaving*, which we will learn more about in this activity.

Why do Karuk and Yurok people prefer to call their prescribed burning practices "cultural burns"?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Hazelnut Stems for Basketweaving
Baskets weaved from *hazelnut stems* have many uses in Karuk and Yurok culture, including (but not limited to) baby cradles, clothing, and food preparation. Here is a hopper basket used to pound acorns. The stems on the left side of the picture are unpeeled, whereas and peeled hazelnut basketry stems are shown on the right. This basket is composed of peeled hazelnut stems along with other materials. 

<img src="Baskets.jpg" alt="drawing" width="300"/>

Skilled basketweavers gather the stems and braid the baskets. Here is a picture of Karuk basketweavers Janet Morehead (left) and Lillian Rentz (right), peeling and evaluating the quality of hazelnut stems they gathered.

<img src="Basketweaver.jpg" alt="drawing" width="400"/>

These basketweavers generously participated in a scientific study that compared the effect of different burn treatments on how hazelnut shrubs produce basket-worthy stems. Additional basketweavers involved in the study include Elizabeth Azzuz, Laverne Glaze, Chook-Chook Hillman, Lisa Hillman, Robert McConnell,Robert McConnell Jr., Kathy McCovey, Bertha Peters, Maggie Peters, Verna Reece, Margo Robbins, Brittany Souza, and Bill Tripp.

# Comparing Burn Treatments for Basket-Worthy Stem Production

We will look at some data from this study [Effects of understory fire management treatments on California Hazelnut, an ecocultural resource of the Karuk and Yurok Indians in the Pacific Northwest](https://www.sciencedirect.com/science/article/pii/S0378112719306826) published in *Forest Ecology and Managment* in 2019. 

Skilled basketweavers (and scientists trained by the skilled basketweavers) collected basket-worthy stems from many different hazelnut shrubs that each received a "treatment". One of the "treatments" was a "broadcast burn" which is the type of cultural burn that the Yurok and Karuk people prefer to use as part of their basketweaving traditions. In the "control," the shrub was left untouched. Stems were gathered before treatment and then again after a full growing season since the treatment. 

The Study Location (yellow star in the intersection of Karuk and Yurok Ancentral Lands) is shown on the map below.

<img src="AncestralTerritories.jpg" alt="drawing" width="700"/>

 Let's have a look at a table of the data...

```python
baskets_and_burning = Table.read_table('Basketweaving_and_Burning.csv')
```

```python
baskets_and_burning
```

**Question 4:** Use table methods to determine how many shrubs recieved "Broadcast" treatment and how many recieved the "Control" treatment?
(*Hint:* `treatment_counts` should be a two column table with each unique treatment and it's corresponding count.)

```python
treatment_counts = ...
treatment_counts
```

```python
grader.check("q1_4")
```

Let's plot the data to visualize it. We will plot the control treatment data using a slightly larger marker than the broadcast burn treatment data in case those data points end up falling on top of each other. We will also plot the data points with some transparency in case the data points within a treatment class fall on top of themselves.

```python
bb_control = baskets_and_burning.where('Treatment', are.equal_to('Control'))
bb_burn = baskets_and_burning.where('Treatment', are.equal_to('Broadcast'))
plots.plot([0,40],[0,40],'k-',linewidth=1)
plots.scatter(bb_control.column('PreBasketStems'),bb_control.column('PostBasketStems'),s=50,alpha=0.7,label='Control')
plots.scatter(bb_burn.column('PreBasketStems'),bb_burn.column('PostBasketStems'),s=30,alpha=0.7,label='Burn')
plots.xlabel('PreBasketStems')
plots.ylabel('PostBasketStems')
plots.legend(loc='upper right')
```

Use the plot above to answer the following questions.

<!-- BEGIN QUESTION -->

**Question 5:** Explain what is plotted on the *x*-axis?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 6:** Explain what is plotted on the *y*-axis?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 7:** What do the *blue* dots represent?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 8:** What do the *red* dots represent?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 9:** What does the *black line* represent?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 10:** Take a look at the two burned shrubs that started out with the most basket-worthy stems. (Red dots with the two largest x-values.) How many basket-worthy stems grew on those shrubs after the burn treatment? What do you think happened to those shrubs?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 11:** Are there any blue dots above the black line? How does this compare to the red dots? What does this mean?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

If you are curious about which datapoints fell on top of each other in the above plot, you can count them using the group function. First we look at the control data:

```python
bb_control.group(['PreBasketStems','PostBasketStems'])
```

And then we can look at the same thing in the broadcast burn data:

```python
bb_burn.group(['PreBasketStems','PostBasketStems'])
```

But perhaps it makes more sense to look at how many basket-worthy stems were on the shrubs after treatment compared to how many basket-worthy stems they started out with.

```python
baskets_and_burning = baskets_and_burning.with_column('BasketStemDiff',baskets_and_burning.column('PostBasketStems')-baskets_and_burning.column('PreBasketStems'))
stemdiff_and_treatment = baskets_and_burning.select('BasketStemDiff', 'Treatment')
```

```python
baskets_and_burning
```

Now we have added a new column showing the difference in basket-worthy stems before and after treatment. Now let's use that new column to make a histogram, where we separate out the data by treatment. We have use some transparency because the data from each treatment sometimes overlap.

```python
stemdiff_and_treatment.hist('BasketStemDiff', group='Treatment')
```

<!-- BEGIN QUESTION -->

**Question 12:** What does a negative number on the x-axis mean?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 13:** According to this sample, which treatment seems best for creating more basket-worthy stems?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

We may wonder whether the difference in the distributions we're looking at is just due to chance or if the treatment has an effect of producing more basket-worthy stems.

We can try to answer this question through an A/B test. 

<!-- BEGIN QUESTION -->

**Question 14:** State the *null hypothesis*, *alternative hypothesis* and *test statistic*.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 15:** Fill in the code below to compute the observed test statistic.

```python
means_table = ...
means_table
```

```python
means = ...
observed_difference = ...
observed_difference
```

```python
grader.check("q1_15")
```

**Question 16:** Create a function to compute tests statistics. *Hint:* Running `difference_of_means(baskets_and_burning, 'Treatment')` should return your test statistic.

```python
def difference_of_means(table, group_label):
    '''Takes: name of table and column label of group-label variable
    Returns: Difference of means of the two groups'''
    
    #table with the two relevant columns
    reduced = ... 
    
    # table containing group means
    means_table = ...
    
    # array of group means
    means = ...
    
    return ...

    
    
    
    
```

```python
grader.check("q1_16")
```

**Question 17:** Fill in the code below to perform your A/B test. 

```python
def one_simulated_difference(table, group_label):
    '''Takes: name of table and column label of group-label variable, 
    Returns: Difference of means of the two groups after shuffling labels'''
    
    # array of shuffled labels
    shuffled_labels = ...
    
    # table of numerical variable and shuffled labels
    shuffled_table = ...
    
    return ...  
    
differences = make_array()

for i in np.arange(2500):
    new_difference = ...
    differences = ...    
    

    
    


```

```python
grader.check("q1_17")
```

Run the following cell to visualize your results.

```python
Table().with_column('Difference Between Group Means', differences).hist()
print('Observed Difference:', observed_difference)
plots.title('Prediction Under the Null Hypothesis');
plots.scatter(observed_difference, 0, color='red', s=40, zorder=3)
```

<!-- BEGIN QUESTION -->

**Question 18:** In your own words, describe the effect of cultural burning on basket-worthy stem growth, as supported by this test.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# How things are now...

Despite the now widespread acknowledgement of the wisdom of cultural burning, these practices are restricted due USDA Forest Service funding and staff reductions, threatened and endangered wildlife being protected in certain seasons, burn restrictions during major wildfire events, and air quality regulations that constrain available burn days. 

Basketweavers also say the departure of supportive USDA Forest Service managers created major set-backs, because they had to re-establish lines of communication with new staff and inform them about their basketry materials and fire treatment
needs. So even with the federal governments increased fuel reduction treatments throughout this region basketweavers and Tribal members often were not informed about the schedules or locations of understory mechanical treatments and broadcast burns, and would miss opportunities to gather hazelnut stems. 

Efforts within the USDA Forest Service known as the ‘Shared Stewardship’ initiative seek to address several of these constraints.

In the meantime basketweavers have been coming up with their own creative solutions to stimulate basket-worthy stem growth on hazelnut shrubs. The figure below shows different types of treatments now being used: (A) manual cutting of the shrub; (B) small pile burn on an indivual shrub; (C) propone torching of a shrub (D) the traditional broadcast burn.
<img src="Treatments.jpg" alt="drawing" width="400"/>

All of these treatments increase basket-worthy stem growth, where the ones that use fire are the most effective. You can read more about the details in [the paper](https://www.sciencedirect.com/science/article/pii/S0378112719306826) if you like.


At home, watch [this video](https://www.youtube.com/watch?v=QQfrQyb4yw4) about the *Land Back* movement. If you want to read more, here is a [news article](https://www.npr.org/2020/08/24/899422710/to-manage-wildfire-california-looks-to-what-tribes-have-known-all-along) that also covered this topic. 

## 2. (optional) The Great British Bake Off

If you want more practice with A/B testing, feel free to do this next example related to the famous TV show The Great British Bake Off.

>"The Great British Bake Off (often abbreviated to Bake Off or GBBO) is a British television baking competition, produced by Love Productions, in which a group of amateur bakers compete against each other in a series of rounds, attempting to impress a group of judges with their baking skills" [Wikipedia](https://en.wikipedia.org/wiki/The_Great_British_Bake_Off)

For every week of the competition, the judges assign one contestant the title "Star Baker". Ultimately, one winner is crowned every season. Using this information, we would like to investigate how winning Star Baker awards affects the odds of winning a season of the show.

<!-- BEGIN QUESTION -->

**Question 2.1.** We want to know whether winning more Star Baker awards *causes* a change in likelihood of winning the season.  Why is it not sufficient to compare star baker rates for winners and losers?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Running an Experiment

We are going to run the following hypothesis test to determine the association between winning and number of Star Baker awards. The population we are examining is every contestant from seasons 2 through 11 of GBBO. We are going to use the following null and alternative hypotheses:

**Null hypothesis:** The distribution of Star Baker awards between contestants who won their season and contestants who did not win their season is the same.

**Alternative hypothesis:** Contestants who win their season of the show will win more Star Baker awards on average.

Our alternative hypothesis is related to our suspicion that contestants who win more Star Baker awards are more skilled, so they are more likely to win the season.

<!-- BEGIN QUESTION -->

**Question 2.2.** Should we use an A/B test to test these hypotheses? If yes, what is our "A" group and what is our "B" group?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

Check your answers with your neighbors before you move on to the next section.

The `bakers` table below describes the number of star baker awards each contest won and whether or not they won their season (`1` if they won, `0` if they did not win). The data was manually aggregated from Wikipedia for seasons 2-11 of the show. We randomized the order of rows as to not spoil the outcome of the show.

```python
bakers = Table.read_table("star_bakers.csv")
bakers.show(3)
```

<!-- BEGIN QUESTION -->

**Question 2.3.** Create a new table called `means` that contains the mean number of star baker awards for bakers who did not win (`won==0`) and bakers that did win (`won==1`). The table should have the column names `won` and `star baker awards mean`.

```python
means = ...
means
```

```python
grader.check("q2_3")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.4.** Visualize the distribution of Star Baker awards for winners and non-winners. You should use the bins we provided.

Hint: You will want to use the group argument of `tbl.hist`. In order to produce several overlayed histograms based on unique values in a given column, we can do something like `tbl.hist(..., group=<col_name>, bins=...)`!


```python
useful_bins = np.arange(0, 7)
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.5.** We want to figure out if there is a difference between the distribution of Star Baker awards between winners and non winners. 

What should the test statistic be? Which values of this test statistic support the null, and which values support the alternative?

If you are in lab, confirm your answer with a neighbor before moving on.

Hint: You should think about what measures we use to describe a distribution. 


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.6.** Set `observed_difference` to the observed test statistic using the `means` table. 


```python
observed_difference = ...
observed_difference
```

```python
grader.check("q2_6")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.7.** Given a table like `bakers`, a label column `label_col`, and a values column `val_col`, write a function that calculates the appropriate test statistic.

*Hint:* Make sure that you are taking the directionality of our alternative hypothesis into account.


```python
def find_test_stat(tbl, label_col, val_col):
    ...

find_test_stat(bakers, "won", "star baker awards")
```

```python
grader.check("q2_7")
```

<!-- END QUESTION -->

When we run a simulation for A/B testing, we resample by **shuffling the labels** of the original sample. If the null hypothesis is true and the star baker award distributions are the same, we expect that the difference in mean star baker awards will be not change when `"won"` labels are changed.

<!-- BEGIN QUESTION -->

**Question 2.8.** Write a function `simulate_and_test_statistic` to compute one trial of our A/B test. Your function should run a simulation and return a test statistic.


```python
def simulate_and_test_statistic(tbl, labels_col, values_col):
    ...

simulate_and_test_statistic(bakers, "won", "star baker awards")
```

```python
grader.check("q2_8")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.9.** Simulate 5000 trials of our A/B test and store the test statistics in an array called `differences`.


```python
# This cell might take a couple seconds to run
differences = make_array()

...
                                                 
differences
```

```python
grader.check("q2_9")
```

<!-- END QUESTION -->

Run the cell below to view a histogram of your simulated test statistics plotted with your observed test statistic.

```python
Table().with_column('Difference Between Group Means', differences).hist(bins=20)
plots.scatter(observed_difference, 0, color='red', s=30, zorder=2)
plots.ylim(-0.1, 1.35);
```

<!-- BEGIN QUESTION -->

**Question 2.10.** Find the p-value for your test and assign it to `empirical_p`.


```python
empirical_p = ...
empirical_p
```

```python
grader.check("q2_10")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.11.** Using a 5% P-value cutoff, draw a conclusion about the null and alternative hypotheses. Describe your findings using simple, non-technical language. What does your analysis tell you about the association between star baker awards and winning? What can you claim about causation from your statistical analysis? Call your GSI or AI over and confirm your answer. 


_Type your answer here, replacing this text._

```python
grader.check("q2_10")
```

<!-- END QUESTION -->

## 3. Submission

Gus wanted to say congrats on finishing Lab 7!! 

<img src="gusgus.jpg" alt="drawing" width="600"/>

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to Canvas. 

**It is your responsibility to make sure your work is saved before running the last cell.**

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1_15 (points: None)
```
>>> abs(observed_difference)
13.839537869062902
```


Tests for q1_16 (points: None)
```
>>> abs(difference_of_means(baskets_and_burning, 'Treatment'))
13.839537869062902
```


Tests for q1_17 (points: None)
```
>>> len(differences) == 2500
True
```


Tests for q1_4 (points: None)
```
>>> (treatment_counts.column('count') == make_array(41,19)).all() or (treatment_counts.column('count') == make_array(19,41)).all()
True
```


Tests for q2_10 (points: None)
```
>>> 0 <= empirical_p < 0.05
True
```


Tests for q2_3 (points: None)
```
>>> means.num_rows
2
```
```
>>> np.round(min(means.column("star baker awards mean")), 2) == 0.65
True
```
```
>>> np.round(max(means.column("star baker awards mean")), 2) == 1.5
True
```


Tests for q2_6 (points: None)
```
>>> isinstance(observed_difference, float)
True
```
```
>>> float(round(observed_difference, 3))
0.848
```


Tests for q2_7 (points: None)
```
>>> np.isclose(round(find_test_stat(bakers, "won", "star baker awards"), 3) - 0.848, 0)
True
```


Tests for q2_8 (points: None)
```
>>> test_stat = round(simulate_and_test_statistic(bakers, "won", "star baker awards"), 3)
>>> -2 < test_stat < 2
True
```
```
>>> np.random.seed(1)
>>> test_stat2 = simulate_and_test_statistic(bakers, "won", "star baker awards")
>>> np.round(test_stat2, 3) == -0.023 or np.round(test_stat2, 3) == -0.132
True
```


Tests for q2_9 (points: None)
```
>>> len(differences)
5000
```
```
>>> abs(np.average(differences)) < 0.05 # On average, your test statistic should be close to 0
True
```
```
>>> all(differences == differences.item(0)) == False # Make sure all of the test statistics are different
True
```
